<problem display_name="webGLDemo">
<script type="loncapa/python">
<![CDATA[
]]>
</script>
<p>
The shapes below can be selected (yellow) or unselected (cyan).
Clicking on them repeatedly will cycle through these two states.
</p>
<p>
If the cone is selected (and not the cube), a correct answer will be
generated after pressing "Check". Clicking on either "Check" or "Save"
will register the current state.
</p>
<customresponse cfn="vglcfn">
  <jsinput gradefn="WebGLDemo.getGrade"
    get_statefn="WebGLDemo.getState"
    set_statefn="WebGLDemo.setState"
    width="400"
    height="400"
    html_file="https://studio.edx.org/c4x/edX/DemoX/asset/webGLDemo.html"
    sop="false"/>
</customresponse>
</problem>

